<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Student-Activity:-Python's-dict-and-friends" data-toc-modified-id="Student-Activity:-Python's-dict-and-friends-1">Student Activity: Python's dict and friends</a></span><ul class="toc-item"><li><span><a href="#1)-Count-the-occurrence-letters-and-get-summary-statistics" data-toc-modified-id="1)-Count-the-occurrence-letters-and-get-summary-statistics-1.1">1) Count the occurrence letters and get summary statistics</a></span></li><li><span><a href="#2)-Calculate-the-sum-of-clicks-per-day" data-toc-modified-id="2)-Calculate-the-sum-of-clicks-per-day-1.2">2) Calculate the sum of clicks per day</a></span></li><li><span><a href="#3)-Change-the-type-of-the-values-in-a-dict" data-toc-modified-id="3)-Change-the-type-of-the-values-in-a-dict-1.3">3) Change the type of the values in a <code>dict</code></a></span></li><li><span><a href="#4)-Convert-data-into-a-ontology" data-toc-modified-id="4)-Convert-data-into-a-ontology-1.4">4) Convert data into a ontology</a></span></li></ul></li></ul></div>

<center><h1>Student Activity: Python's dict and friends</h1></center>

1) Count the occurrence letters and get summary statistics
----

Import `Counter` and count the occurrence of each letter in 

`'abracadabra'`

From the `Counter` instance:

1. Slice out the most frequent letter.
1. Find number of unique letters (cardinality). It will be the same as `set('abracadabra')`.
1. Find total number of letters. It will be the same as `len('abracadabra')`

```python

```

2) Calculate the sum of clicks per day
-----

Given log data with datestamp and number clicks, create an aggregate statistic of total count by date.

<!-- grade_id: cell-5de32fc3defb0c76, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-fa7af86c508faf5f, points: 3, grade: true, solution: false -->
```python
# Test code

logs = [  ('2009-09-02', 11),
          ('2009-09-02', 3),
          ('2009-09-03', 10),
          ('2009-09-03', 4),
          ('2009-09-03', 22),
          ('2009-09-06', 33),
          ('2009-09-02', 13), # Logs are not always in order!
          ('2009-09-03', 15),
       ]

assert sum_clicks_per_day(logs) == (('2009-09-02', 27), 
                                    ('2009-09-03', 51), 
                                    ('2009-09-06', 33))
```

3) Change the type of the values in a `dict`
-----

During data a processing, you noticed that the values of a dictionary are being stored as `str`. They should be `int`. 

Convert all type of each item in values to `int`.

```python
original =  {'A': ['1', '2', '3', '4'],
             'B': ['5', '6', '7', '8'],
             'C': ['9', '10', '11', '12'],}
```

```python

```

```python
# Test code

assert updated == {'A': [1, 2, 3, 4], 
                   'B': [5, 6, 7, 8], 
                   'C': [9, 10, 11, 12]}
```

4) Convert data into a ontology
----

Ontology is a collection of various [taxonomies](https://en.wikipedia.org/wiki/Taxonomy) that can be used to describe a domain of knowledge along with the relationships among them. Ontology is organized metadata of associative relationships between entities.

Ontologies are common in Natural Language Processing (NLP), see [WordNet](https://wordnet.princeton.edu/), and ecommerce, see Amazon.com search.

Given [semantic triple](https://en.wikipedia.org/wiki/Semantic_triple) of subject–predicate–object, group them into a dictionary that captures the categorical relationships. The dictionary should only contain parent(key) and child (value) relationships, ignore all other relationships.

```python
triples = [  ["Ulta", "hasCategory", "Men"],
             ["Ulta", "hasCategory", "Skin Care"],
             ["Ulta", "hasCategory", "Gifts"],
             ["Men", "hasCategory", "Shaving"],
             ["Men", "hasCategory", "Cologne"],
             ["Shaving", "hasCategory", "Aftershave"],
             ["Aftershave", "Contains", "Alcohol"],
             ["Shaving", "hasCategory", "Beard Care"],
             ["Shaving", "hasCategory", "Shaving Cream & Razors"],
             ["Gifts", "hasCategory", "Nail Gifts"],
             ["Gifts", "hasCategory", "Skin Gifts"],
             ["Gifts", "hasCategory", "ULTA Gifts"],
             ["Gifts", "hasCategory", "Makeup Gifts"],
             ["Makeup Gifts", "IsCouponEligible", "False"],
             ["Ulta", "hasCategory", "Tools & Brushes"],
          ]

```

<center><img src="images/ontology.png" width="75%"/></center>

```python
from collections import defaultdict # 👈 hint
```

<!-- grade_id: cell-2ade1e17d92a2155, points: 0, grade: false, solution: true -->
```python

    
```

<!-- grade_id: cell-b007d895d9a078c3, points: 4, grade: true, solution: false -->
```python
# Test code

assert make_ontology(triples) == {   
                                     'Ulta':    {'Gifts', 'Men', 'Skin Care', 'Tools & Brushes'},
                                     'Men':     {'Cologne', 'Shaving'},
                                     'Shaving': {'Aftershave', 'Beard Care', 'Shaving Cream & Razors'},
                                     'Gifts':   {'Makeup Gifts', 'Nail Gifts', 'Skin Gifts', 'ULTA Gifts'},
                                   }

```

<br>
<br> 
<br>

----